Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    # YOUR CODE HERE
    #initated the sum, start_index,end_index and sum_max = 0
    sum = 0
    start_index = 0
    end_index = 0
    sum_max =0
    #for loop i will go up and scan through every entry until it reaches the last number
    for i in range(len(A)):
        sum = 0 #here so that after each iteration of i, the sum will be updated to 0.
        '''
         another loop that runs based on the update of i. i = 0, there will be (len(A)- i )* iteration of j. 
         The sum will be from the first element to the last. 
         Once i = 1, it will calculate the sum of A[1] to last number. 
         The loop continually does that.
        '''
        for j in range(i,len(A)):  
            sum = sum+A[j]
            #if the sum is larger than previous sum_max, sum_max will be updated. Also, it updates the start_index and
            #end_index based the i,j. 
            if sum > sum_max:
                sum_max = sum
                start_index = i
                end_index = j
    return sum_max,start_index, end_index #return the result. 
#raise NotImplementedError() 
A=[-2, -3, 4, -1, -2, 1,8]
print(bruteforce_max_subarray(A))

```

Output:
```
(10, 2, 6)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Yes, I have to re-run the entire maximum array gain because the largest number now is 8. However, between 4 and 8, there are  negative numbers. As 8 is the largest and 4 is the second-largest number, we need to calculate the sum between those numbers. Therefore, we can not conclude whether highest sum is 8 but instead try to re-run the array.

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->


## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    sum = 0
    start_index = 0
    end_index = 0
    sum_max =0
    #for loop i will go up and scan through every entry until it reaches the last number
    for i in range(len(A)):
        sum = 0 #here so that after each iteration of i, the sum will be updated to 0.
        '''
         another loop that runs based on the update of i. i = 0, there will be (len(A)- i )* iteration of j. 
         The sum will be from the first element to the last. 
         Once i = 1, it will calculate the sum of A[1] to last number. 
         The loop continually does that.
        '''
        for j in range(i,len(A)):  
            sum = sum+A[j]
            #if the sum is larger than previous sum_max, sum_max will be updated. Also, it updates the start_index and
            #end_index based the i,j. 
            if sum > sum_max:
                sum_max = sum
                start_index = i
                end_index = j
    return sum_max,start_index, end_index

def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    if len(x)<=1:
        return x
    else:
        #mx equals the maximum subarray of array X excluding its last element.
        mx = bruteforce_max_subarray(x[0:-1])
        return bruteforce_max_subarray(mx)

        
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-18-2256e7bbc067> in <module>
      1 B = [-2, -3, 4, -1, -2, 1, 8]
----> 2 assert(incremental_max_subarray(B, 4) == 10)
      3 assert(incremental_max_subarray(B[:1], 0) == B[0])

AssertionError:
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE