Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **35 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - The rod-cutting problem

The pre-class work for this session will focus on the rod cutting problem. Recall that the rod cutting problem takes as an input the length $n$ of a rodand, a table of prices $p_i$ for $i = 1,2,... n$, and one needs to determine the maximum revenue $r$ obtainable by cutting up the rod and selling the pieces. 

## Part A. 

You will implement in Python two solutions to the rod cutting problem, namely:

## Question 1 [time estimate: 5 minutes] 
A recursive top-down solution to the rod cutting problem. Please complete the cut_rod function below:


<!-- grade_id: cell-d85ce9860c9ea1e5, points: 0, grade: false, solution: true -->
```python
import math
def cut_rod(p,n):
    """
    A recursive top-down solution to the rod cutting problem as described in 
    Cormen et. al. (pp 363) This calculates the maximum revenue that could be 
    earned by cutting up a rod of length n.
    
    Parameters
    ----------
    p : list of floats
        The prices of rods of different lengths. p[i] gives the dollars
        of revenue the company earns selling a rod of length i+1.
    n : int
        Length of the rod
    
    Returns
    -------
    q : float
        The optimal revenue
        
    """
    if n == 0:
        return 0 #if the rod has length 0, then the price would be 0.
    q = -float('inf') #intialize q as sentinel value
    for i in range(n):
        q = max(q,p[i]+cut_rod(p,n-i-1))
    return q
    
    
```

<!-- grade_id: cell-0288a4047be5088d, points: 1, grade: true, solution: false -->
```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1

#results from textbook
r = [0,1,5,8,10,13,17,18,22,25,30]
for i in range(len(r)):
    assert(cut_rod(p,i)==r[i])
```

## Question 2 [time estimate: 10 minutes]

An optimal rod cutting solution using dynamic programming (see figure [here](https://drive.google.com/open?id=1nu9gETKX4KJCHZDi17fZLQtkDVig81Zk) for inspiration). Please complete the following two functions: 


<!-- grade_id: cell-eb49e8516ebf80aa, points: 0, grade: false, solution: true -->
```python
import math
def max(x,y):
    if x >y:
        return x
    return y
def extended_bottom_up_cut_rod(p,n):
    """
    Implements a bottom-up dynamic programming approach to the rod cutting problem.
    Here, "extended" means the function is geared in a way amenable to reconstructing
    an optimal solution, on top of the returned optimal value. See Cormen et al.,
    p. 269 for the implementation details.
    
    Parameters
    ----------
    p : list of floats
        The prices of rods of different lengths. p[i] gives the dollars
        of revenue the company earns selling a rod of length i+1.
    n : int
        Length of the rod
    
    Returns
    -------
    r : list of floats
        The maximum revenues. r[i] gives the maximum revenue for a rod
        of length i. As such:
        
            >>> r[0] = 0
            >>> len(r) == n + 1
            
    s : list of ints
        The optimal sizes of the first piece to cut off. Also make sure 
        that:
            
            >>> s[0] = 0
            >>> len(s) == n + 1
            
    """
    r = [0 for i in range(n+1)] #creating the revenue list
    s = [0 for i in range(n+1)] #list for collecting the optimal sizes to cut off
    for j in range(1, n+1): #iterating n times
        q = -float('inf') #initializing the optimal revenue
        for i in range(j): 
            if q < p[i]+r[j-i-1]: #checking if the revenue is smaller than the expression
                q = p[i]+r[j-i-1] #if yes, then reinitialize the optimal revenue
                s[j] = i+1 #collecting the cut sizes
        r[j] = q #collecting to a list of of maximum revenues at each step
    return r,s #returning both lists
```

<!-- grade_id: cell-1d52c9519ccb3502, points: 0, grade: false, solution: true -->
```python
def print_cut_rod_solution(p,n):
    """
    Gives a solution to the rod cutting problem of size n. 
    
    Parameters
    ----------
    p : list of floats
        The prices of rods of different lengths. p[i] gives the revenue 
        (in USD, for example) the company earns selling a rod of length i+1
    n : int
        Length of the rod
    
    Returns
    -------
    sol : a list of ints
        Indicates how to cut the rod. Cutting the rod with the lengths
        given in sol gives the optimal revenue.
        
            >>> print_cut_rod_solution(p,0) == []
            
    """
    r,s = extended_bottom_up_cut_rod(p,n) 
    while n > 0:
        print(s[n]) #printing the optimal cut sizes 
        n = n - s[n]
```

<!-- grade_id: cell-6575aa2769e0eab1, points: 1, grade: true, solution: false -->
```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1
# Result of R and S from textbook:
R = [0,1,5,8,10,13,17,18,22,25,30]
S = [0,1,2,3,2,2,6,1,2,3,10]
# Test:
r, s = extended_bottom_up_cut_rod(p,10)
assert(r==R)
assert(s==S)
```

## Part B - Experiments

## Question 1 [time estimate: 1 minute]

Use the function below to generate a list of prices of length n=20 and assign that list to a new variable, `P`. You MUST use this list for questions 2 and 3 below.


```python
import numpy as np
def generate_price_list(n):
    """
    Generates a price list of length n
    
    Parameters
    ----------
    n : int
        Length of the list, must be positive
    
    Returns
    -------
    p : list
        The ordered price list for each rod cut
        
    """
    p = [1]
    for i in range(1,n):
        np.random.seed(0)
        p.append(np.random.randint(p[i-1]+1, i*3+1))
    return p
P = generate_price_list(20)
print(P)
```

Output:
```
[1, 2, 3, 8, 9, 14, 15, 20, 21, 26, 27, 32, 33, 38, 39, 44, 45, 50, 51, 56]
```

## Question 2 [time estimate: 5 minutes]

Time (using the time library) the implementation of `cut_rod` for different rod lengths for values of up to $n=20.$ Present your results in a plot. 


<!-- grade_id: cell-e094e913d8bc139f, points: 0, grade: true, solution: true -->
```python
import time
import matplotlib.pyplot as plt
attempts = 15
leng = 20
times = []
for i in range(leng+1):
    total_time = 0
    for attempt in range(attempts):
        start = time.time()
        cut_rod(P,n)
        time += time.time() - start
    times.append(time/attempts)
    
plt.ylabel("Time in miliseconds")
plt.xlabel("Lengths")
plt.plot([number for number in range(lengths+1)],times, color='magenta')
plt.show()
         
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-30-e4721bcf34dc> in <module>
      8     for attempt in range(attempts):
      9         start = time.time()
---> 10         cut_rod(P,n)
     11         time += time.time() - start
     12     times.append(time/attempts)

NameError: name 'n' is not defined
```

## Question 3 [time estimate: 2 minutes]

Time (using the time library) the implementation of `extended_bottom_up_cut_rod` for different rod lengths for values of up to $n=20$. Add the curve corresponding to this algorithm to the previous plot.

<!-- grade_id: cell-210aa3d5a2743bd8, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 4 [time estimate: 2 minutes]

Contrast both curves. What can you conclude?

<!-- grade_id: cell-0e62cc5b19840488, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Part C - After completing your experiments above, answer the following questions:

## Question 1 [time estimate: 5 minutes]
Estimate how big $n$ has to be before the non-dynamic-programming version will take a million years to finish. Call this value N.

<!-- grade_id: cell-a11f4fd6a03e2cf5, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2 [time estimate: 5 minutes]
Estimate (or time) how long it takes to evaluate `extended_bottom_up_cut_rod(p,N)`, where `N` is the value you got from the previous question. 

<!-- grade_id: cell-89ed2dd6f3310c96, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## [Optional]  Question 3 [time estimate: 5 minutes]
Do you notice anything peculiar about the solution for large n?

<!-- grade_id: cell-fd4c9c91a42f178a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE