Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan "
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq
#def add_to_median_heap(minh, maxh, elem):
 #   if not maxh or minh:
A= [1,2,8,2,4,9,10]
heapq.heappush(A,1)
print(A)
heapq.heappop(A)


    
    
    #raise NotImplementedError()
```

Output:
```
[1, 1, 8, 2, 4, 9, 10, 2]
1
```

```python
a = []
#if a has element
if a:
    print("C")
#if a is empty
if not a:
    print("D")
```

Output:
```
D
```

```python
import heapq
def add_to_median_heap(minh,maxh,elem):
    #if maxh is empty:
    if not maxh:
        #push an element into maxh.
        heapq.heappush(maxh,elem)
        print(maxh)
        return maxh
    #takes the largest element of maxh.
    m = max(maxh)
    print(m)
    #compares elem value vs maximum value of maxh.
    if elem >= m:
        #if elem is smaller than maximum value of maxh,
        #push the element into minh because maxh could not take any value larger than its root.
        heapq.heappush(minh,elem)
        #balance the heap in case one heap has two more elements. 
        if len(minh) == len(maxh)+2:
            heapq.heappush(maxh,-heapq.heappop(minh))
            
    else:
        heapq.heappush(maxh,elem)
        if len(maxh) == len(minh)+2:
            heapq.heappush(minh,heapq.heappop(maxh))
    return maxh,minh

```

```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    if not maxh :
        #because heapq implements a min heap
        heapq.heappush(maxh,-elem)
    elif maxh: 
    #compare new element with top of the heap
        if elem > -maxh[0]: 
            heapq.heappush(minh, elem)
        else:
            heapq.heappush(maxh, -elem)
    #balance the heap
        if len(minh) == len(maxh)+2:
            j = heapq.heappop(minh)
            heapq.heappush(maxh,-j)
        elif len(maxh) == len(minh)+2:
            j = -heapq.heappop(maxh)
            heapq.heappush(minh,j)
```

## Question 2 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    if not maxh and not minh:
        return None
    else: 
        median = (minh[0]+maxh[0])/2
        if len(minh) > len(maxh): # if minh has more elements than maxh, then the median is in minh
            median = minh[0]
        else: # if maxh has more elements than minh, then the median is in maxh
            median = -maxh[0]
    return median # return the median
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
<ipython-input-40-2a7cc854325b> in <module>
      3 for a in range(1,100,2):
      4     add_to_median_heap(minh, maxh, a)
----> 5     print(median(minh, maxh))

<ipython-input-38-47a4905b4af5> in median(minh, maxh)
      3         return None
      4     else:
----> 5         median = (minh[0]+maxh[0])/2
      6         if len(minh) > len(maxh): # if minh has more elements than maxh, then the median is in minh
      7             median = minh[0]

IndexError: list index out of range
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The worst case complexity to building a heap is O(n).
Therefore,  using add_to_median_heap does not cost more than that because there are no recursions and it only compares and uses push and poll, so the worst case complexity is also O(n)

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
Median would only use comparisons, therefore the worst case can happen would be O(1).

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->


## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def partition(A,p,r): # defining the partition function from quicksort 
    x = A[r]
    i = p
    for j in range(p, r):
        if A[j] <= x:
            A[i], A[j] = A[j], A[i]
            i += 1

    A[i], A[r] = A[r], A[i]
    return i

def quickselect(lst, s, r, k):
    if (k > 0 and k <= r - s + 1): 
        x = partition(lst , s, r) # partition in the side of k

        if (x - s == k - 1): # if found, return the value of it
            return lst[x] 

        if (x - s > k - 1): 
            return quickselect(lst, s, x - 1, k) # call the function recursively
        return quickselect(lst, x + 1, r, k - x + s - 1)
    
def qselect(lst,k):
    s = 0
    r = len(lst) - 1 # define s and r to use the previous function
    quickselect(lst, s, r, k) #call the function
    return(lst[k])

```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-43-253afcee666d> in <module>
      6 for a in range(100):
      7     lst2.append(qselect(lst1, a))
----> 8 assert(lst2 == list(range(100)))

AssertionError:
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-41-de9182726f6f> in <module>
      4 random.shuffle(lst)
      5 for a in range(100):
----> 6     print(qselect(lst, a))

NameError: name 'qselect' is not defined
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
The recurrence relation happens when I call the function of quickselect after partioning the list recursively until I find the position of k.



## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
Best case: O(n).
It is due to the fact that it only solves the problem in one side which has k.



## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE