Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **48 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Hiring problem and random variables

## Part A. The Hire-Assistant Problem.

Imagine that you need to hire a new assistant. Every day an agency sends a new assistant for you to interview. If the assistant is better than your current assistant, then you fire your current assistant and you hire the better assistant. You may assume that assistant quality is uniformly distributed between 0 and 1.

## Question 1 [time estimate: 5 minutes]
Write a function, named hire_assistant, that takes applicants (a list of the numbers that represent the level of qualification of the applicants; the higher the number, the better qualified), and returns the number hires if the applicants are presented in the exact same order as the input list applicants. Note that your function should not randomize anything (or else it would be called a randomized algorithm).

<!-- grade_id: cell-d011f5f4707fe41a, points: 0, grade: false, solution: true -->
```python
def hire_assistant(applicants):
    """
    This function should output the number of assistants hired.
    
    Parameters
    ----------
    applicants : list
        A list of the numbers that represent the level of qualification of 
        the applicants; the higher the number, the better qualified.
        
    Returns
    -------
    hires : int
        Number of assistants hired.
    
    """
    timeshire= 0
    best_candidate = 0
    for i in range(len(applicants)):
        if applicants[i]>best_candidate:
            best_candidate = applicants[i]
            applicants[i]
            timeshire+=1
    return timeshire
A=[1,2,3,4,5,6]
print(hire_assistant(A))
    # YOUR CODE HERE
#raise NotImplementedError()
```

Output:
```
6
```

<!-- grade_id: cell-66778b97ad66f71e, points: 1, grade: true, solution: false -->
```python
assert(hire_assistant([1])==1)
assert(hire_assistant([-1, -2, -3, -4])==1)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-23-2dddc1aba31f> in <module>
      1 assert(hire_assistant([1])==1)
----> 2 assert(hire_assistant([-1, -2, -3, -4])==1)

AssertionError:
```

## Question 2 [time estimate: 5 minutes]

Assuming the applicants are presented in a random order, write a function that receives the number of applicants as input and returns the average number of assistants hired.

**N.B.:** Don’t forget to run the simulation several times for each given number of applicants to better estimate the number of hires (please refer to task 3 of the Study Guide).

<!-- grade_id: cell-b223520ca72942a0, points: 0, grade: true, solution: true -->
```python
import random
def experimental_hires(N):
    timeshire= 0
    best_candidate = 0
    for i in range(len(applicants)):
        if applicants[i]>best_candidate:
            best_candidate = applicants[i]
            applicants[i]
            timeshire+=1
    return (timeshire//len(A))
A= random.randrange(1,10,1)
print(experimental_hires(A))
    
    
#raise NotImplementedError()
```

Output:
```
<ipython-input-25-f7e76136c81a> in <module>
     10     return (timeshire//len(A))
     11 A= random.randrange(1,10,1)
---> 12 print(experimental_hires(A))
     13 
     14 

<ipython-input-25-f7e76136c81a> in experimental_hires(N)
      3     timeshire= 0
      4     best_candidate = 0
----> 5     for i in range(len(applicants)):
      6         if applicants[i]>best_candidate:
      7             best_candidate = applicants[i]

NameError: name 'applicants' is not defined
```

## Question 3 [time estimtate: 10 minutes]

Use the function below, `analytical_hires(N)`, which returns the analytical expected number of hires, given the number of applicants, along with the function you created in question 2 to create a graph with two curves such that:
* The x-axis shows the total number of applicants (make sure label the x-axis)
* The y-axis shows the average number of hires (make sure label the y-axis)
* The graph contains two curves;
    * Curve 1: the theoretical performance estimates computed calls to the function `analytical_hires`.
    * Curve 2: the simulated or experimental estimates using the function you created in question 2.


```python
def analytical_hires(N):
    """
    This function returns the analytical expected number of
    hires if there are N applicants
    
    Parameters
    ----------
    N : int
        Number of applicants
        
    Returns
    -------
    hires : float
        Average number of assistants hired
    
    Note
    ----
    From the textbook, we know that the analytical result is
    
        1 + 1/2 + 1/3 + ... + 1/N
        
    """
    hires = 0
    for n in range(N):
        hires += 1/(n+1)
    return hires
```

<!-- grade_id: cell-f9c07920c069ce20, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 4 [time estimate: 10 minutes]

Plot a graph with the x-axis showing the total number of applicants and the y-axis showing the probability that exactly one assistant is hired.

<!-- grade_id: cell-d3fe1b7d6d175ad7, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## [Optional] Question 5 [time estimate: 15 minutes]
Assume that an assistant is able to perform an amount of work each day that is equal to their “quality”. You have a total amount of work M that needs to be accomplished. Your costs are as follows:
* X = daily salary for the assistant,
* Y = fee to the employment agency,
* Z = retrenchment fee for the old assistant.

Try to formulate an optimal stopping rule (i.e., at what point should one stop requesting new potential hires from the agency?) Make any necessary assumptions to ensure the problem is well-formulated.


<!-- grade_id: cell-af2f0291eced6982, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. The Hat Check Problem.

There is a coat check at a party, where an attendant stores everyone’s hat while they attend the party. The attendant receives the N hats from everyone attending (all attendees come with a hat). Unfortunately, the coat check attendant forgets which hat belongs to whom. Rather than admitting a mistake, the attendant simply returns random hats back to the party-goers. 

What is the average number of correct hats returned? Here are some guiding questions to help you to simulate this problem. 

## Question 1 [time estimate: 3 minutes]
Knowing that everyone’s hats are unique and every guest has a hat, do you need to generate a random sample in a similar way as what you did for the hiring assistant problem? 

<!-- grade_id: cell-e786799fc4eb1499, points: 0, grade: true, solution: true -->


## Question 2 [time estimate: 3 minutes]

Which of the following commands do you think is the Pythonic way to implement that? 
```
import numpy as np
n = 100 #the number of party attendants
```
**Command 1.**
```
hat_list = [np.random.integers(0,n) for i in range(n)]
```
**Command 2.**
```
hat_list = list(range(n)) 
np.random.shuffle(hat_list) 
```
**Command 3.**
```
hat_list = np.random.sample(n)
```

<!-- grade_id: cell-b8da78e72c1c0738, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3 [time estimate: 5 minutes]

Now write a function `hat_check(N)` that has: 
* Input: N the number of party attendants. 
* Output: the number of hats correctly returned despite the fact that hats are randomly handed back to the guests.

You should use the command you picked for question 2. 

<!-- grade_id: cell-c8499aeb1b1d76c7, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 4 [time estimate: 7 minutes]

Plot a curve with the x-axis showing the total number of party attendants and the y-axis showing the average number of hats correctly returned. As always, remember to run several trials. 

<!-- grade_id: cell-597031ea2a5a512a, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## [Optional] Question 5 [time estimate: 15 minutes]

As $N$ tends to infinity, the number of correct hats returned tends towards a well-known statistical distribution. State the distribution with all its parameters. Plot several samples using your code. Does the empirical distribution match your theoretical prediction?

<!-- grade_id: cell-32fe26c1d99fdd2a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE