Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **34 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - The Delta Median problem and Randomization

## Question 1 [time estimate: 16 minutes]

### Question 1a [time estimate: 7 minutes]

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
import random
def issort(A):
    for i in range(1,len(A)):
        if A[i]<A[i-1]:
            return False
    return True
def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Parameters
    ----------
    A : list
        A list of numerical values
        
    Returns
    -------
    A : list
        Sorted list
        
    """
    while issort(A)== False:
        random.shuffle(A)
    return A
A= [3,6,4,5,2,5]
print(worst_sort(A))
```

Output:
```
[2, 3, 4, 5, 5, 6]
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b [time estimate: 3 minutes]
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
The best case complexity of the algorithm is when the list A is already sorted, which we will need to execute issort(A) once in the function worst_sort. The running time of issort(A) would equals O(n), because the for loop going through the whole list (n steps) and  making the step equals to n.Therefore, the best case complexity is $O( n)$

### Question 1c [time estimate: 3 minutes]

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
The average case complexity is $O(N!*N) because N! equals the number of ways that the list would be sorted and N equals the number of steps each time issort(A) is executed. 

### Question 1d [time estimate: 3 minutes]

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
When the size list is small because the average case is already every big. 

### [Optional] Question 1e [time estimate: 5 minutes]

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2 [time estimate: 18 minutes]

Approximate median finder. Given a list and $\tilde{\delta}$ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the $(50-\tilde{\delta}/2)$% and $(50+\tilde{\delta}/2)$% percentiles, where $\tilde{\delta}=100\cdot \delta$. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a [time estimate: 5 minutes]
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is 
    within (50-delta/2)% and(50+delta/2)% percentiles.
    
    X% percentile of a list is defined to be the smallest number in 
    the list that is as large as at least X% numbers in the list. 
    
    Parameters
    ----------
    A : list
        A list of numerical values
    num 
        The number of interest
    delta : float
        The error, lies between 0 and 0.5

    Returns
    -------
    isApproxMed : bool
        A boolean value indicating whether num is whithin the bound
    
    """
    #if the len of list A equals 1 and the first item equals the number, return true. otherwise, return false.
    if len(A) == 1:
        if A[0] == num:
            return True
        return False
    #calculate the percentile by using the location. 
    else:
        location = 0
        for i in A:
            if num > i:
                location += 1
        percentile = location / len(A)
        if percentile > 0.5 - (delta / 2) and percentile < 0.5 + (delta / 2):
            return True
        else:
            return False
    
    # YOUR CODE HERE
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0, 1], 0.4, .25) == True)
assert(check_approx_med([0, 1], 0, .25) == False)
assert(check_approx_med([0], 0, .25) == True)
```

### Question 2b [time estimate: 5 minutes]
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 
    50-delta/2% and 50+delta/2% percentiles
    
    Paramters
    ---------
    A : list
        A list of numerical values
    Delta : int
        The error, lies between 0 and 0.5
        
    Returns
    -------
    num
        The approximated median if it is found within 100//delta trials
        (this is one possible maximum number of retries. While you are 
        encouraged to play around with another limits, the submitted work 
        must use this number)
    None
        Finding failed, if nothing is found within 100//delta trials
   
   
    Note
    ----
    The 100//delta is chosen here as it is the expected number of trials for 
    the first successful finding to occur (using geometric distribution)
    
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

### Question 2c [time estimate: 3 minutes]

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with $\delta$ or $N$ (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### Question 2d [time estimate: 5 minutes]
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE