Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **44 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Randomized quicksort

## Part A. Median-of-3 partitioning quicksort 

## Question 1 [time estimate: 10 minutes]

Read through the following Python code. What does each function (i.e., median, qsort, randomized_qsort, test_qsort) do? Comment in details each function. 


```python
#import two libraries called timeit and random
import timeit
import random


eps = 1e-16
#length of the array we sort
N = 10000
locations = [0.0, 0.5, 1.0 - eps]

#define a function that tries to find a median by comparing three inputs x1,x2,x3 and return the largest second element.
def median(x1, x2, x3):
    if (x1 < x2 < x3) or (x3 < x2 < x1):
        return x2
    elif (x1 < x3 < x2) or (x2 < x3 < x1):
        return x3
    else:
        return x1



def qsort(lst):
    #create a list containing a tuple that includes first and length of the list
    indices = [(0, len(lst))]
     #if the indeces exit:
    while indices:
        #return the list containing fr,to- which can be the first and last element of indeces. 
        (frm, to) = indices.pop()
        #check if the range equals 0, it will return to the loop. 
        if frm == to:
            continue
        
        # Find the partition:
        # N is the length of the list
        N = to - frm
        # this creates three numbers, first, middle and last indexes because n in locations has the value of 0.0, 0.5 and 1.0-eps, 
        #so when multiplying with N(length of the list)+ frm, we get the actual index of three values. 
        inds = [frm + int(N * n) for n in locations]
        #based on inds- indexes of three numbers, it creates a list containing all three numbers. 
        values = [lst[ind] for ind in inds]
        #find the median by using the defined function above. 
        partition = median(*values)
        
        # Split into lists:
        #creates a list that value smaller than median
        lower = [a for a in lst[frm:to] if a < partition]
        #creates a list that its value larger than median
        upper = [a for a in lst[frm:to] if a > partition]
        #counts the total of elements equals to median, each element would add 1 to the total sum
        counts = sum([1 for a in lst[frm:to] if a == partition])
        #index of last element of left-sided list 
        ind1 = frm + len(lower)
        #index of the first element of the right-sided list 
        ind2 = ind1 + counts

        # Push back into correct place:
        lst[frm:ind1] = lower
        lst[ind1:ind2] = [partition] * counts
        lst[ind2:to] = upper

        # Enqueue other locations
        indices.append((frm, ind1))
        indices.append((ind2, to))
    return lst


def randomized_quicksort():
    lst = [i for i in range(N)]
    random.shuffle(lst)
    return qsort(lst)


def test_quicksort():
    lst = randomized_quicksort()
    assert (lst == [i for i in range(N)])


# Is our algorithm correct
test_quicksort()

# How fast is our algorithm
print(timeit.timeit(randomized_quicksort, number=1))
```

Output:
```
0.0516948999993474
```

## Question 2 [time estimate: 3 minutes]

What are the main differences between the `randomized_quicksort` in the code and $RANDOMIZED-QUICKSORT$ in Cormen et al., besides that the partition of `randomized_quicksort` uses a median of 3 as a pivot?

<!-- grade_id: cell-4a3cd727ccac7404, points: 0, grade: true, solution: true -->
The difference between the randomized_quicksort here and in cormen is that the code in Cormen uses sorting in place while here it creates upper/lower lists, increasing the time complexity. 

## Question 3 [time estimate: 10 minutes]

What is the time complexity of this `randomized_qsort`? Time the algorithm on lists of various lengths, each list being a list of the first $n$ consecutive positive integers. Produce a graph with list lengths on the x axis and running time on the y axis. As always, don’t forget to time the algorithm several times for each list’s length and then average the results. 

<!-- grade_id: cell-e0e1dac71ac7feb6, points: 0, grade: true, solution: true -->
```python
import time
import matplotlib.pyplot as plt
a =[]
b=[]
#Using for loop to check the average runtime for using quicksort to sort array length range from 1 to 1000:
for i in range(1,1000):
    time =[]
    for k in range(10):
        lst = [l for l in range(i)]
        random.shuffle(lst)
        t_start= time.perf_counter()
        qsort(lst)
        t_end = time.perf_counter() - t_start
        # append the runtime into t
        time.append(t_end)
    a.append(i)
    b.append(sum(t_end)/10)
plt.scatter(a,b, s=1)
plt.xlabel('Length of the array')
plt.ylabel('Second')
plt.show()
    
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-15-8ca1096320d0> in <module>
      9         lst = [l for l in range(i)]
     10         random.shuffle(lst)
---> 11         t_start= time.perf_counter()
     12         qsort(lst)
     13         t_end = time.perf_counter() - t_start

AttributeError: 'list' object has no attribute 'perf_counter'
```

## Question 4.

### Question 4a [time estimate: 7 minutes]

Change the `qsort()` function in a way that you **don’t** separate the items that are equal to the partition. 


<!-- grade_id: cell-a9d1f063c0340b14, points: 0, grade: false, solution: true -->
```python

#import two libraries called timeit and random
import timeit
import random


eps = 1e-16
#length of the array we sort
N = 10000
locations = [0.0, 0.5, 1.0 - eps]

#define a function that tries to find a median by comparing three inputs x1,x2,x3 and return the largest second element.
def median(x1, x2, x3):
    if (x1 < x2 < x3) or (x3 < x2 < x1):
        return x2
    elif (x1 < x3 < x2) or (x2 < x3 < x1):
        return x3
    else:
        return x1



def qsort(lst):
    #create a list containing a tuple that includes first and length of the list
    indices = [(0, len(lst))]
     #if the indeces exit:
    while indices:
        #return the list containing fr,to- which can be the first and last element of indeces. 
        (frm, to) = indices.pop()
        if frm <=2:
            if lst[to-1] < lst[frm]:
            #swap element if latter element is smaller than previous element.
                lst[frm], lst[to-1] = lst[to-1], lst[frm]
            continue
        
        # Find the partition:
        # N is the length of the list
        N = to - frm
        # this creates three numbers, first, middle and last indexes because n in locations has the value of 0.0, 0.5 and 1.0-eps, 
        #so when multiplying with N(length of the list)+ frm, we get the actual index of three values. 
        inds = [frm + int(N * n) for n in locations]
        #based on inds- indexes of three numbers, it creates a list containing all three numbers. 
        values = [lst[ind] for ind in inds]
        #find the median by using the defined function above. 
        partition = median(*values)
        
        # Split into lists:
        #creates a list that value smaller than or equal to median
        lower = [a for a in lst[frm:to] if a <= partition]
        #creates a list that its value larger than median
        upper = [a for a in lst[frm:to] if a > partition]
        #counts the total of elements equals to median, each element would add 1 to the total sum
        counts = sum([1 for a in lst[frm:to] if a == partition])
        #index of last element of left-sided list 
        ind1 = frm + len(lower)
        # Push back into correct place:
        lst[frm:ind1] = lower
        lst[ind1:to] = upper

        # Enqueue other locations
        indices.append((frm, ind1))
        indices.append((ind1, to))
    return lst


```

<!-- grade_id: cell-2c0cbd296d612f85, points: 1, grade: true, solution: false -->
```python
assert(qsort([4,2,1])==[1,2,4])
assert(qsort([0])==[0])
```

### Question 4b [time estimate: 3 minutes]

Now time the algorithm on the same inputs you have used in question 3, adding one more line in the previous graph you have produced. 

<!-- grade_id: cell-31ee807cec9ce8bf, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Question 5.

### Question 5a [time estimate: 3 minutes]

Remove the median-of-3 partitioning, and just use the first element in the array. 

<!-- grade_id: cell-4daf36021c15eaf0, points: 0, grade: false, solution: true -->
```python
def qsort(lst):
    #create a list containing a tuple that includes first and length of the list
    indices = [(0, len(lst))]
     #if the indeces exit:
    while indices:
        #return the list containing fr,to- which can be the first and last element of indeces. 
        (frm, to) = indices.pop()
        if frm <=2:
            if lst[to-1] < lst[frm]:
            #swap element if latter element is smaller than previous element.
                lst[frm], lst[to-1] = lst[to-1], lst[frm]
            continue
        
        # Find the partition:
        partition = lst[frm]
        
        # Split into lists:
        #creates a list that value smaller than or equal to median
        lower = [a for a in lst[frm:to] if a <= partition]
        #creates a list that its value larger than median
        upper = [a for a in lst[frm:to] if a > partition]
        #counts the total of elements equals to median, each element would add 1 to the total sum
        counts = sum([1 for a in lst[frm:to] if a == partition])
        #index of last element of left-sided list 
        ind1 = frm + len(lower)
        # Push back into correct place:
        lst[frm:ind1] = lower
        lst[ind1:to] = upper

        # Enqueue other locations
        indices.append((frm, ind1))
        indices.append((ind1, to))
    return lst
```

<!-- grade_id: cell-97473a9e0d12e745, points: 1, grade: true, solution: false -->
```python
assert(qsort([4,2,1])==[1,2,4])
assert(qsort([0])==[0])
```

### Question 5b [time estimate: 3 minutes]

Does this change the running time of your algorithm? Justify your response with a graph. 



<!-- grade_id: cell-1f3a6df29d324853, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Part B. Recursive quicksort [time estimate: 5 minutes]

One main difference between the quicksort algorithms in Cormen et al. and the implementation in the code above is that quick sort (in the code in this notebook) is not recursive, while $QUICKSORT$ in Cormen et al. is. Given the limitation of Python so that it can only make 500 recursive calls, estimate the maximum size of the list that can be sorted by Python if a recursive quicksort is to be used. Explicitly state all assumptions you make in getting to an answer.


<!-- grade_id: cell-4af5aab4ad1a7225, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE