Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **37 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan Nguyen"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Asymptotic notations; maximum-subarray problem

## Question 1 (Exercise 3.1-3 of Cormen, et al. ) [time estimate: 5 minutes]
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is true but meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
As you use big-O notation, it indicates the upper tight bound - the maximum running time that one algorithm takes. Here, it says at least, which indicates the lower bound of the running time of the algorithm. Therefore, the claim does not provide any information about upper bound. 

## Question 2 (Exercise 3.1-4 of Cormen, et al. ) [time estimate: 5 minutes]

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$? Why?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
The first is correct as $2^{n+1}$ equals $2*2^n$. $O(2^n)$ can be interpreted c*2^n. Here c = 2. Therefore, the statement is correct. 

## Question 3 [time estimate: 20 minutes]
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    # YOUR CODE HERE
    #initated the sum, start_index,end_index and sum_max = 0
    sum = 0
    start_index = 0
    end_index = 0
    sum_max =0
    #for loop i will go up and scan through every entry until it reaches the last number
    for i in range(len(A)):
        sum = 0 #here so that after each iteration of i, the sum will be updated to 0.
        '''
         another loop that runs based on the update of i. i = 0, there will be (len(A)- i )* iteration of j. 
         The sum will be from the first element to the last. 
         Once i = 1, it will calculate the sum of A[1] to last number. 
         The loop continually does that.
        '''
        for j in range(i,len(A)):  
            sum = sum+A[j]
            #if the sum is larger than previous sum_max, sum_max will be updated. Also, it updates the start_index and
            #end_index based the i,j. 
            if sum > sum_max:
                sum_max = sum
                start_index = i
                end_index = j
    return sum_max,start_index, end_index #return the result. 
#raise NotImplementedError() 
A=[9,-1,-9,2,-11,18,29,-4]
print(bruteforce_max_subarray(A))
```

Output:
```
(47, 5, 6)
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

## Question 4 [time estimate: 2 minutes]
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HE
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
print(bruteforce_max_subarray(A))
#raise NotImplementedError()
```

Output:
```
(43, 7, 10)
```

## Question 5. Asymptotic notation [time estimate: 5 minutes]
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Running time 
--- | --- 
Insertion sort | 
Selection sort |  
Bubble sort |  
Finding maximum subarray |  

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->

Algorithm | Running time 
--- | --- 
Insertion sort | $O(n^2)$ 
Selection sort |  $O(n^2)$
Bubble sort | $O(n^2)$ 
Finding maximum subarray |  $O(n^2)$

## [Optional] Question 6 [time estimate: 15 minutes] 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """
    Implements brute-force minimum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the minimum subarray
        - the end index of the minimum subarray
        - the value of the minimum subarray
        
    """
    sum = 0
    start_index = 0
    end_index = 0
    sum_min =0
    #for loop i will go up and scan through every entry until it reaches the last number
    for i in range(len(A)):
        sum = 0 #here so that after each iteration of i, the sum will be updated to 0.
        '''
         another loop that runs based on the update of i. i = 0, there will be (len(A)- i )* iteration of j. 
         The sum will be from the first element to the last. 
         Once i = 1, it will calculate the sum of A[1] to last number. 
         The loop continually does that.
        '''
        for j in range(i,len(A)):  
            sum = sum+A[j]
            #if the sum is larger than previous sum_max, sum_max will be updated. Also, it updates the start_index and
            #end_index based the i,j. 
            if sum < sum_min:
                sum_min = sum
                start_index = i
                end_index = j
    return sum_min,start_index, end_index #return the result. 
#raise NotImplementedError() 
A=[9,-1,-9,2,-11,18,29,-4]
print(bruteforce_min_subarray(A))
    # YOUR CODE HERE
# raise NotImplementedError()
```

Output:
```
(-19, 1, 4)
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```