Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **63 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan "
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Computational Applications of Randomized Algorithms

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson Heaps and Heapsort, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1 [time estimate: 10 minutes]
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq
#def add_to_median_heap(minh, maxh, elem):
 #   if not maxh or minh:
A= [1,2,8,2,4,9,10]
heapq.heappush(A,1)
print(A)
heapq.heappop(A)


    
    
    #raise NotImplementedError()
```

Output:
```
[1, 1, 8, 2, 4, 9, 10, 2]
1
```

```python
a = []
#if a has element
if a:
    print("C")
#if a is empty
if not a:
    print("D")
```

Output:
```
D
```

```python
import heapq
#define a function that takes in three different parameters,
def add_to_median_heap(minh, maxh, elem):
    #if the maxheap and minheap is empty:
    #push the element into minheap.
    #return.
    if not maxh and not minh:
        heapq.heappush(minh, elem)
        return
    #if the maxheap is empty:
    if not maxh:
        #compare maxheap with minheap[0]- root of minheap.
        if elem > minh[0]:
            #push the element into maxheap, at the same time multiply by 1 so that the element popping out from 
            #min heap would become the smallest element when adding into in the maxheap.
            #we pop out the recent element from maxheap because we need to update the minheap- the half side
            #containing element larger than median.
            heapq.heappush(maxh, -heapq.heappop(minh))
            print(maxh)
            #then after popping out the previous element, we have new element added into the minheap.
            heapq.heappush(minh, elem)
        #if the element is smaller than minh[0]    
        else:
            #we push new element into the maxheap by multiplying it with -1,
            #to preserve the structure of the maxheap.
            heapq.heappush(maxh, -elem)
        return
    #if length of maxheap and minheap equals:
    if len(maxh) == len(minh):
        #if element is smaller than the root *-1 of maxheap
        #maxh[0] is the largest element, therefore, it means that elem is smaller than smallest 
        #element of maxh.
        #push -element into maxheap 
        if elem < -maxh[0]:
            heapq.heappush(maxh, -elem)
        else:
            heapq.heappush(minh, elem)
        #two other scenarios executing the operations based on comparison of length of maxheap and minheap
    elif len(maxh) > len(minh):
        if elem < -maxh[0]:
            heapq.heappush(minh, -heapq.heappop(maxh))
            heapq.heappush(maxh, -elem)
        else:
            heapq.heappush(minh, elem)
    else:
        if elem > minh[0]:
            heapq.heappush(maxh, -heapq.heappop(minh))
            heapq.heappush(minh, elem)
        else:
            heapq.heappush(maxh, -elem)
```

```python
a =[]
if not a:
    print("A")
```

Output:
```
A
```

## Question 2 [time estimate: 5 minutes]
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
"""
    Objective
    The code below finds the median number of a list using max and min heaps
    
    Input
    array: A representation of the min heap holding all the numbers larger than the median
    array: A representation of the max heap holding all the numbers larger than the median
    
    Output:
    int: The median number in the sequence of numbers
"""


def median(minh, maxh):
    if len(maxh) == len(minh):
        return float((-maxh[0] + minh[0]) / 2)
    elif len(maxh) > len(minh):
        return float(-maxh[0])
    else:
        return float(minh[0])
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1.0
2.0
3.0
4.0
5.0
6.0
7.0
8.0
9.0
10.0
11.0
12.0
13.0
14.0
15.0
16.0
17.0
18.0
19.0
20.0
21.0
22.0
23.0
24.0
25.0
26.0
27.0
28.0
29.0
30.0
31.0
32.0
33.0
34.0
35.0
36.0
37.0
38.0
39.0
40.0
41.0
42.0
43.0
44.0
45.0
46.0
47.0
48.0
49.0
50.0
```

## Question 4 [time estimate: 3 minutes]
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The worst case complexity to building a heap is O(n).
Therefore,  using add_to_median_heap does not cost more than that because there are no recursions and it only compares and uses push and poll, so the worst case complexity is also O(n)

## Question 5 [time estimate: 3 minutes] 
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
Median would only use comparisons, therefore the worst case can happen would be O(1).

## Question 6 [time estimate: 4 minutes]

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
As we use sorting approach, we need to sort the list in order to find the median in constant time. In constrast, when using the median approach, we reduce the time complexity because our priority is no longer sorting the list but finding the median. 
Therefore, partitioning the items to a max heap and min heap and since the median is located at the root, it sorts in linear time and lookup is in constant time respectively.

## [Optional] Question 7 [time estimate: 15 minutes]

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
#raise NotImplementedError()
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1 [time estimate: 10 minutes]

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
"""
    Function
    Qselect procedure by finding and displaying the kth smallest element in a list of numbers
    
    Input:
    array: The list of numbers from which the number is selected
    int: The position of the number to be selected
    
    Output:
    int: The k-th smallest number in the list.
"""

def qselect(lst,k):
    def select(lst, l, r, index):
        # base case
        if r == l:
            return lst[l]
        # choose random pivot
        pivot_index = random.randint(l, r)

        # move pivot to beginning of list
        lst[l], lst[pivot_index] = lst[pivot_index], lst[l]

        # partition
        i = l
        for j in range(l+1, r+1):
            if lst[j] < lst[l]:
                i += 1
                lst[i], lst[j] = lst[j], lst[i]

        # move pivot to correct location
        lst[i], lst[l] = lst[l], lst[i]

        # recursively partition one side only
        if index == i:
            return lst[i]
        elif index < i:
            return select(lst, l, i-1, index)
        else:
            return select(lst, i+1, r, index)

    if lst is None or len(lst) < 1:
        return None
    if k < 0 or k > len(lst) - 1:
        raise IndexError()
    return select(lst, 0, len(lst) - 1, k)

```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

## Question 2 [time estimate: 1 minute]

Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.

```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))
```

Output:
```
0
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
```

## Question 3 [time estimate: 5 minutes]

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
The recurrence relation happens when I call the function of quickselect after partioning the list recursively until I find the position of k.



## Question 4 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
Best case: O(n).
It is due to the fact that it only solves the problem in one side which has k.



## Question 5 [time estimate: 3 minutes]

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
Quickselect's worst-case is $O(n^2)$

```python

```