Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **50 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Tuan"
COLLABORATORS = "Google pleases to accompany us :D "
```

---

# CS110 Pre-class Work - Analyzing elementary sorting algorithms



## Question 1 [time estimate: 10 minutes]
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
It helps to break from the loop and decrease the value of i by 1. 
Previous, we stored the value of key and we kept the value so that in case the value of A[i] > key, we swap the value of A[j] with the value of A[i]. 

Then, as our previous A[j] is kept as key, we update the value of A[i] with that. 


For instance, we have an array of 46,16, 39. Key is 16. At line 6, the value of array is [46,46,39]. 

It then decreases the value of i =0, so that the loop breaks and the value of i index will be incremented by 1 and updated, which creates an array of [16,46,39]. If we do not have line 7, the array will remain the same. 

## Question 2 [time estimate: 15 minutes]
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertion_sort(A):
    for j in range(2, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertion_sort(A):
    for j in range(1, len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i+1] = key
    return A
    raise NotImplementedError()
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertion_sort([0]) == [0])
assert(insertion_sort([-1,1]) == [-1,1])
assert(insertion_sort([1,-1]) == [-1,1])
assert(insertion_sort([1,6,3,6]) == [1,3,6,6])
assert(insertion_sort([7, 3, 5, 4, 6, 4]) == [3, 4, 4, 5, 6, 7])
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
When the code does not have error when running. Also, the code works when when we change the range, as we considered A[..,i]- sorted and A[j,...]- unsorted, which i starts at 1 is essential for comparison of i and j.

## Question 3 [time estimate: 15 minutes]
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertion_sort`,  your new Python function, `bubble_sort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubble_sort(A):
    for i in range(1,len(A)-1):
        for j in range(len(A)-i-1):
            if arr[j] > arr[j+1] : 
                arr[j], arr[j+1] = arr[j+1], arr[j]
array = [64, 34, 25, 12, 22, 11, 90] 
  
bubble_sort(array) 
  
print ("Sorted array is:") 
for i in range(len(array)): 
    print ("%d" %array[i]),  
#raise NotImplementedError()
```

Output:
```
Sorted array is:
64
34
25
12
22
11
90
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
```python
assert(bubble_sort([1,6,3,6]) == [1,3,6,6])
assert(bubble_sort([7, 3, 5, 4, 6, 4]) == [3, 4, 4, 5, 6, 7])
assert(bubble_sort([-1,1,15,10,15,16]) == [-1,1,10,10,15,15,16])

```

## Question 4 [time estimate: 10 minutes]
The following pseudocode corresponds to a third sorting algorithm covered in this lesson, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free by providing at least three test cases. `selection_sort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selection_sort(A):
    n = len(A)
    for i in range(1,n-1):
        minidx = i
        for j in range (i+1,len(A)):
            if A[j]<A[minidx]:
                minidx = j
        A[i], A[minidx] = A[minidx], A[i] 

array = [64, 34, 25, 12, 22, 11, 90] 
  
selection_sort(array) 
for i in range(len(array)): 
    print ("%d" %array[i])
  
print (selection_sort(array))
    #raise NotImplementedError()
```

Output:
```
64
11
12
22
25
34
90
None
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**